---
title: Test
description: Test
free_preview: true
---

## Sample exercise

```yaml
type: NormalExercise
key: d52caa96ad
lang: python
xp: 100
skills: 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
