---
title: Test
description: Test
free_preview: true
---

## Sample exercise

```yaml
type: NormalExercise
key: b29fb4c5bb
lang: python
xp: 100
skills: 2
```

Change1Change1Change1Change1Change1Change12

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
